Negative regulation of miRNAs sorting in EVs: the RNA-binding protein PCBP2 impairs SYNCRIP-mediated miRNAs EVs loading

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

In this study, Marocco and colleages perform a deep characterization of the complex molecular mechanism guiding the recognition of a particular CELLmotif previously identified in hepatocytes in another publication. Having miR-155-3p with or without this CELLmotif as initial focus, authors identify 21 proteins differentially binding to these two miRNA versions. From these, they decided to focus on PCBP2. They elegantly demonstrate PCBP2 binding to miR-155-3p WT version but not to CELLmotif-mutated version. miR-155-3p contains a hEXOmotif identified in a different report, whose recognition is largely mediated by another RNA-binding protein called SYNCRIP. Interestingly, mutation of the hEXOmotif contained in miR-155-3p did not only blunt SYNCRIP binding, but also PCBP2 binding despite the maintenance of the CELLmotif. This indicates that somehow SYNCRIP binding is a pre-requisite for PCBP2 binding. EMSA assay confirms that SYNCRIP is necessary for PCBP2 binding to miR-155-3p, while PCBP2 is not needed for SYNCRIP binding. Then authors aim to extend these finding to other miRNAs containing both motifs. For that, they perform a small-RNA-Seq of EVs released from cells knockdown for PCBP2 versus control cells, identifying a subset of miRNAs whose expression either increases or decreases. The assumption is that those miRNAs containing PCBP2-binding CELLmotif should now be less retained in the cell and go more to extracellular vesicles, thus reflecting a higher EV expression. The specific subset of miRNAs having both the CELLmotif and hEXOmotif (9 miRNAs) whose expressions increase in EVs due to PCBP2 reduction is also affected by knocking-down SYNCRIP in the sense that reduction of SYNCRIP leads to lower EV sorting. Further experiments confirm that PCBP2 and SYNCRIP bind to these 9 miRNAs and that knocking down SYNCRIP impairs their EV sorting.

In the revised manuscript, the authors have addressed most of my concerns and questions. I believe the new experiments provide stronger support for their claims. My only remaining concern is the lack of clarity in the replicates for the EMSA experiment. The one shown in the manuscript is clear; however, the other three replicates hardly show that knocking down SYNCRIP has an effect on PCBP2 binding. Even worse is the fact that these replicates do not support at all that PCBP2 silencing has no effect on SYNCRIP binding, as the bands for those types of samples are, in most of the cases, not visible. I think the authors should work on repeating a couple of times EMSA experiment.

We thank this Reviewer for having appreciated the novelty and the robustness of our data. In accordance with the Reviewer’s concern, we repeated the EMSA assay, specifically to address the PCBP2-independent SYNCRIP binding. In Author response image 1, we report the new EMSA replicates (top), the quantification of each signal (bottom) and the mean of EMSA signals relative to the three independent experiments (right). We hope that the new evidence will meet the required standards.

Reviewer #2 (Public review):

Summary:

The author of this manuscript aimed to uncover the mechanisms behind miRNA retention within cells. They identified PCBP2 as a crucial factor in this process, revealing a novel role for RNAbinding proteins. Additionally, the study discovered that SYNCRIP is essential for PCBP2's function, demonstrating the cooperative interaction between these two proteins. This research not only sheds light on the intricate dynamics of miRNA retention but also emphasizes the importance of protein interactions in regulating miRNA behavior within cells.

Strengths:

This paper makes important progress in understanding how miRNAs are kept inside cells. It identifies PCBP2 as a key player in this process, showing a new role for proteins that bind RNA. The study also finds that SYNCRIP is needed for PCBP2 to work, highlighting how these proteins work together. These discoveries not only improve our knowledge of miRNA behavior but also suggest new ways to develop treatments by controlling miRNA locations to influence cell communication in diseases. The use of liver cell models and thorough experiments ensures the results are reliable and show their potential for RNA-based therapies

Weaknesses:

The manuscript is well-structured and presents compelling data, but I noticed a few minor corrections that could further enhance its clarity:

Figure References: In the response to Reviewer 1, the comment states, "It's not Panel C, it's Panel A of Figure 1"-this should be cross-checked for consistency.

Supplementary Figure 2 is labeled as "Panel A"-please verify if additional panels (B, C, etc.) are intended.

Western Blot Quality: The Alix WB shows some background noise. A repeat with optimized conditions (or inclusion of a cleaner replicate) would strengthen the data. Adding statistical analysis for all WBs would also reinforce robustness.

These are relatively small refinements, and the manuscript is already in excellent shape. With these adjustments, it will be even stronger.

We deeply thank this Reviewer for having considered this new version of the manuscript and for having described its shape as excellent. In order to address the Reviewer’s concerns, we crosschecked the consistency of the described figures’ panels described in the text accordingly. Regarding the qualitative analysis of EV markers, we repeated the western blot analysis with optimized conditions as suggested and included the new panel (Author response image 2) in the supplementary figure 2, allowing to appreciate the signal relative to ALIX expression.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

Careful reading is required to rectify typo errors.

We thank the Reviewer for this suggestion. We amended the text to rectify typo errors.

Changes from the earlier version to the revised version

Abstract

Abstract

While it is accepted that Extracellular Vesicles (EVs)-mediated transfer of microRNAs contributes to intercellular communication, the knowledge about molecular mechanisms controlling the selective and dynamic miRNA-loading in EVs is still limited to few specific RNA-binding proteins interacting with sequence determinants. Moreover, although mutagenesis analysis demonstrated the presence/function of specific intracellular retention motifs, the interacting protein/s remained unknown.

Here, PCBP2 was identified as a direct interactor of an intracellular retention motif: CLIP coupled to RNA pull down and proteomic analysis demonstrated that it binds to miRNAs embedding this motif and mutagenesis proved the binding specificity. Notably, PCBP2 binding requires SYNCRIP, a previously characterized miRNA EV-loader as indicated by SYNCRIP knock-down.

SYNCRIP and PCBP2 may contemporarily bind to miRNAs as demonstrated by EMSA assays and PCBP2 knock-down causes EV-loading of intracellular microRNAs.

This evidence highlights that multiple proteins/miRNA interactions govern miRNA compartmentalization and identifies PCBP2 as a dominant inhibitor of SYNCRIP function.

Introduction

Extracellular Vesicles (EVs) budding from cellular plasma membrane and differing for size and origin (e.g. exosomes, of 30-150 nm diameter and microvesicles, which diameter is ≥ 50 nm) (1) are uptaken by neighboring as well as by distant recipient cells thus transferring their specific informational cargo of molecules. Indeed, EVs are now considered a well assessed route for donor cells to promote changes in gene expression and cell behavior of recipient cells in several pathophysiological processes (2, 3). Notably, concerning miRNAs, the EV-cargo does not simply reflect the cell of origin content, but rather is defined by dynamic and selective cell-specific loading mechanisms (4–8). The existence of multiple players controlling miRNA compartmentalization matches with the hypothesis of a multiprotein machinery that dynamically governs not only the EV sorting but also, conceivably, the intracellular retention of specific subsets of miRNAs.

Earlier version: At the present, the molecular players controlling the selective partition of miRNAs seems largely uncharacterized although recent evidence highlighted mechanisms of EV-loading depending on sequence-specific RNA-binding proteins (RBPs). On this regard, the first identified was the Heterogeneous Nuclear Ribonucleoprotein A2B1 (hnRNPA2B1), recognizing miRNAs containing the EXO motif (G/U,G,A/G/C,G/C) and guiding their inclusion in EVs secreted by human primary T-lymphocytes (8). A functional role was successively attributed to the Synaptotagmin-binding Cytoplasmic RNA-Interacting Protein (SYNCRIP or hnRNPQ) that, by binding to miRNAs embedding a short hEXO motif (G/A/U,G/U/A,G/A/U,C/A/G,U/A,G/C) through its NURR domain and its RRMs recognizing the 5’ of the motif (7, 9), guides their inclusion in EVs secreted by hepatocytes. Notably, hEXO was shown to increase the EV/cell ratio of miRNAs as demonstrated by the insertion of this sequence in a cell-retained miRNA (7, 9) thus paving the way for further manipulative perspectives in the growing field of RNA-based therapies (10–12).

Revised version: At present, the molecular players controlling the selective partition of miRNAs seems largely uncharacterized although recent evidence highlighted mechanisms of EV-loading depending on sequence-specific RNA-binding proteins (RBPs). On this regard, the first identified was the Heterogeneous Nuclear Ribonucleoprotein A2B1 (hnRNPA2B1), recognizing miRNAs containing the EXO motif (G/U,G,A/G/C,G/C) and guiding their inclusion in EVs secreted by human primary T-lymphocytes (8). A functional role was successively attributed to the Synaptotagmin-binding Cytoplasmic RNA-Interacting Protein (SYNCRIP or hnRNPQ) that, by binding to miRNAs embedding a short hEXO motif (G/A/U,G/U/A,G/A/U,C/A/G,U/A,G/C) through its NURR domain and its RRMs recognizing the 5’ of the motif (7, 9), guides their inclusion in EVs secreted by hepatocytes. Notably, hEXO was shown to increase the EV/cell ratio of miRNAs as demonstrated by the insertion of this sequence in a cell-retained miRNA (7, 9) thus paving the way for further manipulative perspectives in the growing field of RNA-based therapies (10–12).

More recently, similar functional evidence gathered by the chimeric approach has been extended to newly identified motifs, causing miRNAs export or intracellular retention (5). While for these export motifs two RNA-binding proteins (ALYREF and FUS) have been characterized (5) for the retention motifs, defined as “CELL” (5), as for the previously described “CL” motifs (8) the identification of interacting protein/s remains unaddressed.

We here aimed at the investigation of molecular players responsible for miRNAs intracellular retention. Overall, we gathered evidence on the interactions among several CELL-embedding miRNAs and two RBPs: SYNCRIP and the multifunctional RNA-binding protein PCBP2. Functionally, PCBP2 promotes cellular retention in SYNCRIP-bound miRNAs and the already described SYNCRIP loading activity appears impeded by PCBP2.

Results

PCBP2 recognizes a CELL motif and has a functional role in intracellular retention of miRNA-155-3p

Aiming to identify RBPs involved in the intracellular retention, proteins from hepatocytes were used in RNA-pull-down by using as specific bait the miRNA-155-3p, selected for the presence of the Core CELL-motif identified in AML12 cells (5) (Figure 1A).

PCBP2 recognizes the CELL motifs and has a functional role in intracellular retention of miRNA-155-3p.

A) Sequences of biotinylated oligos used as bait in pulldown experiment; Core CELL motifs (WT and mutated) are in grey, hEXO motifs (WT and mutated) are underlined. miRNA devoid of CELL motif (no-CELL), miRNA devoid of hEXO (no-hEXO).

Earlier version: B) Volcano plot comparing proteins bound to miRNA-155-3p no-CELL vs WT. Black curves represent the significant threshold at an a false-discovery rate (FDR) of 0.05 and S0 of 0.1. PCBP2 and SYNCRIP proteins are labeled in the plot.

Revised version: B) Volcano plot comparing proteins bound to miRNA-155-3p no-CELL vs WT. Black curves represent the significant threshold at a false-discovery rate (FDR) of 0.05 and S0 of 0.1. PCBP2 and SYNCRIP proteins are labeled in the plot.

C) CLIP of PCBP2 protein in murine hepatocytes. RT-qPCR analysis for miR-155-3p, miR-365-2-5p (CELL motif-devoid) and miR-31-3p (hEXO motif-devoid) is shown as IP/IgG. Data are the mean± SEM of three independent experiments.

D) RNA pull-down with the WT and mutated (no-CELL, no-hEXO) (sequences are reported in A) miR-155-3p followed by western blot for the indicated proteins (HSP90 is used as positive and GAPDH as negative controls respectively). Data are representative of three independent experiments.

E) RNA pull-down with the WT and mutated (+CELL) miR-365-2-5p followed by western blot for the indicated proteins. Data are representative of three independent experiments.

F) RNA pull-down by using the recombinant PCBP2 protein and with WT and mutated miR-155-3p (no-CELL) followed by western blot for PCBP2. Data are representative of three independent experiments.

G) CLIP of SYNCRIP protein in murine hepatocytes. RT-qPCR analysis for miR-155-3p is shown as IP/IgG. Data are the mean± SEM of three independent experiments.

H) (left and middle panels) EV miRNA-155-3p and miR-365-2-5p levels in shCTR and shPCBP2 cells analyzed by RT-qPCR. Data are expressed as ratio of miRNA expression in EVs with respect to the intracellular compartment (shCTR arbitrary value 1). Results are shown as the mean±S.E.M. of three independent experiments.

(right panel) EV miRNA-155-3p levels in shCTR and shSYNCRIP cells analyzed by RT-qPCR. Data are expressed as ratio of miRNA expression in EVs with respect to the intracellular compartment (shCTR arbitrary value 1). Results are shown as the mean±S.E.M. of three independent experiments. Data are considered statistically significant with p< 0.05 (Student’s T test). *: p<0.05; **: p<0.01.

Label-free nLC-MS/MS proteomic analysis allowed to identify miR-155-3p-interacting proteins (Supplementary Data File 1) and Label-free quantification intensities analysis identified 21 proteins enriched in miR-155-3p pulldown with respect to miR-155-3p mutated in the Core CELL motif (miR-155-3p no-CELL) (Figure 1A, Supplementary Data File 2 and Figure 1B). Twelve of them are classified as RNA-binding proteins, with six containing at least one canonical RNA-binding domain (13–16) (Supplementary Data File 2). For three of them, RNA-binding preferences are also known (i.e. PCBP2 prefers CU-rich sequences (15, 17), LARP1 recognizes the CAP and the 5’ top motif in mRNAs (18) and STAU1 binds to double-stranded RNAs (19)).

Earlier version: Among them, PCBP2 interaction with miR-155-3p was confirmed by CLIP assay (Figure 1C, left panel) and RNA pull-down followed by western blot analysis (Figure 1D). A second proof of concept of the requirement of PCBP2/CELL motif interaction is provided by the observation on the CELL motif-devoid miR-365-2-5p (Figure 1C, middle panel); RNA pull-down confirmed the absence of interaction with PCBP2 (Figure 1E). Notably, the inclusion of the Core CELL motif while maintaining SYNCRIP binding, allows PCBP2 interaction with this miRNA. At least in vitro, this binding appears direct and sequence specific as demonstrated by the use of a recombinant protein in RNA pull-down (Figure 1F). The introduction of specific mutations allowed to test the requirement of the CELL motif since its modification (miR-155-3p no-CELL) impairs PCBP2 binding (Figure 1D).

Revised version: Among them, PCBP2 interaction with miR-155-3p was confirmed by CLIP assay (Figure 1C, left panel) and RNA pull-down followed by western blot analysis (Figure 1D). A second proof of concept of the requirement of PCBP2/CELL motif interaction is provided by the observation on the CELL motif-devoid miR-365-2-5p (Figure 1C, middle panel); RNA pull-down confirmed the absence of interaction with PCBP2 (Figure 1E). Notably, the inclusion of the Core CELL motif while maintaining SYNCRIP binding, allows PCBP2 interaction with this miRNA. At least in vitro, this binding appears direct and sequence specific as demonstrated using a recombinant protein in RNA pull-down (Figure 1F). The introduction of specific mutations allowed to test the requirement of the CELL motif since its modification (miR-155-3p no-CELL) impairs PCBP2 binding (Figure 1D).

Notably, the inspection of the miR-155-3p sequence revealed the presence of the previously identified SYNCRIP binding site and both MS/MS analysis (Supplementary Data File 1) and CLIP assay (Figure 1G) confirmed SYNCRIP binding to this miRNA independently of the CELL motif mutation (Figure 1D).

Surprisingly, the permutation of the downstream two nucleotides removing SYNCRIP-binding motif (miR-155-3p no-hEXO, Figure 1A) impairs PCBP2 binding despite the conservation of the CELL retention motif (Figure 1D).

This suggests a possible SYNCRIP requirement for PCBP2 binding. To test this hypothesis, CLIP assay was performed on miR-31-3p, embedding the sole CELL motif, indicating the absence of PCBP2 binding (Figure 1C, right panel).

Functionally, the role of PCBP2 in miRNA partition was addressed by its silencing (Supplementary Figure 1A and B). Notably, PCBP2 interference (PCBP2 being not expressed in the EVs, Supplementary Figure 2A) enhances miRNA-155-3p loading in EVs with respect to control cell-derived EVs (Figure 1H, left panel). As a further control that miRNAs without the CELL motif are not affected by PCBP2 silencing, the expression levels of miR-365-2-5p (embedding the sole hEXO motif) were analyzed in EVs and cells, and resulted not differentially exported (Figure 1H, middle panel). As expected SYNCRIP silencing (Supplementary Figure 1C-D) (expressed in both intracellular and EV compartment as shown in Supplementary Figure 2A) reduces miR-155-3p export into EVs (Figure 1H, right panel); (for EVs characterization see Supplementary Figure 2 A and B). Moreover, the effect of either PCBP2 or SYNCRIP knockdown on EVs production was ruled out by quantifying the number of EVs per cell in shPCBP2 or shSYNCRIP with respect to the shCTR conditions (Supplementary Figure 2 C).

Overall, these data demonstrated that i) PCBP2 interacts with miRNA-155-3p, as proved by CLIP analysis and RNA pull-down, ii) the interaction is CELL-motif-dependent while an unexpected role for the hEXO motif is also unveiled and iii) PCBP2 favors the intracellular localization of this miRNA.

PCBP2 binding to miR-155-3p is both sequence- and SYNCRIP-dependent

The observation that loading (hEXO) and retention (CELL) motifs are both present in miR-155-3p sequence prompted us to investigate on the hypothesis of a sequence- and SYNCRIP-dependent PCBP2 binding ability.

First, we observed that the two proteins interact each other (Figure 2A) and more interestingly that both RBPs bind to miR-155-3p contemporarily as indicated by the ultra-shift obtained in EMSA assay (Figure 2B).

PCBP2 binding to miR-155-3p is SYNCRIP-dependent

A) Co-immunoprecipitation of PCBP2 and SYNCRIP. Immunoprecipitations with rabbit polyclonal anti-PCBP2, mouse monoclonal anti-SYNCRIP and the relative preimmune IgG were performed on protein extracts from hepatocytes. GAPDH is used as negative control. Immunoblots representative of three independent experiments are shown.

B) Electrophoretic mobility shift assay (EMSA): interactions of miR-155-3p with the indicated protein extracts (shifts) and Abs (anti-SYNCRIP and anti-PCBP2) (supershift) are shown. Ultrashift shown in lane 5 demonstrates concurrent binding of SYNCRIP and PCBP2 to miR-155-3p.

C) Electrophoretic mobility shift assay (EMSA): interactions of miR-155-3p with protein extracts from shCTR (1), shPCBP2 (2) and shSYNCRIP (3) cells (shifts) and Abs (anti-SYNCRIP and anti-PCBP2) (supershift) are shown.

D) CLIP of PCBP2 protein in murine hepatocytes both WT (shCTR) and silenced for SYNCRIP (shSYNCRIP). RT-qPCR analysis for the expression of miR-155-3p is shown as IP/IgG. Data are the mean± SEM of three independent experiments.

Earlier version: To challenge the hypothesis of a SYNCRIP-dependent PCBP2 binding, EMSA assay was performed in PCBP2-silenced and in SYNCRIP-silenced cells (Supplementary Figure 1A, B, C and D respectively). As shown in figure 2C while the PCBP2 silencing does not affect SYNCRIP binding, SYNCRIP silencing impairs also PCBP2 binding. Furthermore, CLIP assay was performed on SYNCRIP-silenced cells; as shown in figure 2D, SYNCRIP silencing impairs PCBP2 binding to miR-155-3p.

Revised version: To challenge the hypothesis of a SYNCRIP-dependent PCBP2 binding, EMSA assay was performed in PCBP2-silenced and in SYNCRIP-silenced cells (Supplementary Figure 1A, B, C and D respectively). As shown in Figure 2C while PCBP2 silencing does not affect SYNCRIP binding, SYNCRIP silencing impairs also PCBP2 binding. Furthermore, CLIP assay was performed on SYNCRIP-silenced cells; as shown in Figure 2D, SYNCRIP silencing impairs PCBP2 binding to miR-155-3p.

This evidence supports the unpredictable mechanism where SYNCRIP binding appears a prerequisite for PCBP2 recruitment. To further confirm and extend this observation, a number of mutants were designed and tested by RNA pull-down for the binding capacity of these two proteins; results indicate that i) mutagenesis of the sole hEXO motif (miR-26b-3p no-hEXO) on miR-26 backbone (bearing both hEXO and CELL motifs), impairs also PCBP2 binding (Figure 3A) and conversely ii) the de novo inclusion of a hEXO motif in miR-31-3p backbone (bearing only CELL-motifs) (miR-31-3p +hEXO) confers a de novo PCBP2 binding ability to this mutant; furthermore, mutation in the CELL motif (miR-31-3p +hEXO no-CELL) impairs PCBP2 binding (Figure 3B).

PCBP2 binding to miR-155-3p is sequence dependent

A) RNA pull-down with the WT and mutated (sequences are reported above) miR-26b-3p followed by western blot for the indicated proteins. Data are representative of three independent experiments.

B) RNA pull-down with the WT and mutated (sequences are reported above) miR-31-3p followed by western blot for the indicated proteins. Data are representative of three independent experiments. A-B) Core CELL motifs (WT and mutated) are in grey, hEXO motifs (WT and mutated) are underlined.

Earlier version: Overall, these data indicate that PCBP2 binding requires both the CELL motif and SYNCRIP binding; in other word, SYNCRIP binding is epistatic to PCBP2 recruitment.

Revised version: Overall, these data indicate that PCBP2 binding requires both the CELL motif and SYNCRIP binding; in other words, SYNCRIP binding is epistatic to PCBP2 recruitment.

PCBP2 functionally dominates on SYNCRIP EV-loading activity on a repertoire of miRNAs embedding CELL and hEXO motifs

To extend the evidence for the role of PCBP2 in miRNA compartmentalization and to confirm its mechanistic role i) PCBP2 and ii) SYNCRIP functional role in miRNA EVs/cell partition was evaluated, and iii) PCBP2, iv) SYNCRIP and v) SYNCRIP-dependent PCBP2 binding were assessed. First, NGS analysis of miRNAs exported in EVs produced by control and PCBP2-silenced murine hepatocytes allowed the selection of further miRNAs differentially loaded in EVs in correlation to PCBP2 (Supplementary Data File 3, Figure 4A, B).

PCBP2 functionally dominates on SYNCRIP EV-loading activity on a repertoire of miRNAs embedding CELL and hEXO motifs

A) Volcano plot comparing miRNAs differentially expressed from NGS data; miRNAs with Log2FC > 1 and Log2FC < −1 and p-value ≤ 0.10 were considered differentially expressed. Downregulated miRNAs in shPCBP2 respect to shCTRL are represented as blue dots, upregulated miRNAs are represented as red dots.

B) Heatmap showing the Log2 fold enrichment (EV/CELL) of mature miRNAs in small extracellular vesicles derived from shCTRL cells versus shPCBP2 cells. miRNAs with Log2FE ≥ 1.0 and p-value ≤ 0.10 were considered to be differentially enriched.

C) List of selected miRNAs embedding CELL and/or hEXO motifs; consensus sequences are highlighted in grey or underlined respectively.

D) EV miRNA levels in shCTR and shPCBP2 cells analyzed by RT-qPCR. Data are expressed as ratio of miRNA expression in EVs with respect to the intracellular compartment (shCTR arbitrary value 1). Results are shown as the mean±S.E.M. of three independent experiments.

Data are considered statistically significant with p< 0.05 (Student’s T test). *: p<0.05; **: p<0.01.

E) Gene Ontology enrichment analysis on validated and predicted targets. The X axis represents the gene ratio, whereas the Y axis represents the enriched GO terms. Colors indicate the statistical significance after multiple test correction, while circle size represents the number of genes associated with each term.

F) Liver-isolated leukocytes cultured either alone (NO EVs) or in the presence of control EVs (shCTR EVs) or shPCBP2 EVs (shPCBP2 EVs) for 24 hours. The expression of IFN-γ was analyzed on CD3+, NK1.1+ CD11b+ and CD11b+ cells by flow cytometry. The dot plots show the percentage of cytokine-positive T cells, NK cells and myeloid cells. T cells and NK cells graphs are representative of 5 independent experiments, while myeloid cells graph is representative of 4 independent experiments. Each symbol represents data obtained from an individual mouse.

Data are considered statistically significant with p< 0.05 (Student’s T test). *: p<0.05; **: p<0.01

Then, the functional role of PCBP2 was assessed by means of qRT-PCR performed on 9 miRNAs expressed in EVs and embedding both CELL and hEXO motifs (see Methods section) in comparison to 4 miRNAs embedding either the CELL (miR-31-3p, miR-195a-5p, miR-421-3p) or the hEXO motif (miR-365-2-5p) (Figure 4C).

Specifically, miR-155-5p, miR-181d-5p, miR-3084-5p, miR-122b-3p, miR-192-5p, miR-26b-3p, miR-31-3p, miR-195a-5p and miR-421-3p have the Core CELL motif (AUUA/G) described by Garcia-Martin and colleagues (5). Other miRNAs (miR-345-3p, miR-23a-5p and miR-214-3p) share the described CELL motif (5) with the most frequent nucleotides, considering also the reported variability. Regarding the hEXO motif described by Santangelo and collaborators (7), the most frequent nucleotides defining the motif sequence have been taken into consideration.

Earlier version: In order to consider a possible impact of PCBP2 on miRNAs steady state level (resulting from variation in transcription and biogenesis), miRNA abundance in EVs and in the intracellular compartment was analyzed by qRT-PCR as EVs/cell ratio. Results indicate that PCBP2 silencing releases a SYNCRIP-dependent loading of miRNAs in the EVs (Figure 4D) thus highlighting a dominant PCBP2 cell-retention function on SYNCRIP-dependent export, evaluated on these miRNAs in Supplementary Figure 3.

Revised version: In order to consider the possible impact of PCBP2 on miRNAs steady state level (resulting from variation in transcription and biogenesis), miRNA abundance in EVs and in the intracellular compartment was analyzed by qRT-PCR as EVs/cell ratio. Results indicate that PCBP2 silencing releases a SYNCRIP-dependent loading of miRNAs in the EVs (Figure 4D) thus highlighting a dominant PCBP2 cell-retention function on SYNCRIP-dependent export, evaluated on these miRNAs in Supplementary Figure 3.

Moreover, Gene Ontology (GO) enrichment analysis of the PCBP2-dependent cell-retained miRNAs targets have been performed (including both validated and predicted targets, respectively obtained from TarBase v9.0 database and DIANA-microT web server). As reported in Figure 4E and in the Supplementary Figure 4, this analysis highlighted several biological pathways collectively influenced by these miRNAs (e.g. establishment of organelle localization, regulation of cell cycle and lymphocyte differentiation).

To preliminarily explore the functional impact of EVs deriving from WT and PCPB2-silenced cells, treatment of leukocytes obtained from C57BL/6J mice livers has been performed.

Earlier version: Specifically, as reported in Figure 4F and Supplementary Figure 5, T cells, NK cells and myeloid cells show a higher amount of IFN-γ production upon shPCBP2 EVs treatment in comparison to the shCTR EVs; this suggests that PCBP2 silencing results into an EV-mediated modulation of the immune response.

Revised version: Specifically, as reported in Figure 4F and Supplementary Figure 5, T cells, NK cells and myeloid cells show a higher amount of IFN-γ production upon shPCBP2 EVs treatment in comparison to the shCTR EVs; this suggests that PCBP2 silencing results in an EV-mediated modulation of the immune response.

With respect to the molecular mechanism, the analysis by CLIP-qPCR assay demonstrates the PCBP2 binding to miRs-345-3p, 23a-5p, 214-3p, 155-5p, 181d-5p, 3084-5p, 122b-3p, 192-5p, 26b-3p that bear both the CELL and hEXO motifs; conversely, the presence of either the sole hEXO (miR-365-2-5p) or the sole CELL (miR-31-3p, miR-195a-5p, miR-421-3p) does not allow PCBP2 binding (Figure 5A). As expected, the presence of hEXO motif alone or in combination with the CELL motif is sufficient for SYNCRIP binding to all the analyzed miRNAs (Figure 5B).

PCBP2 and SYNCRIP bind to several miRNAs embedding CELL and hEXO motif sequences

A) CLIP of PCBP2 protein in murine hepatocytes. RT-qPCR analysis for the indicated miRNAs is shown as IP/IgG for each independent experiment (IgG arbitrary value 1). Data are the mean± SEM of three independent experiments.

Data are considered statistically significant with p< 0.05 (Student’s T test). *: p<0.05; **: p<0.01.

B) CLIP of SYNCRIP protein in murine hepatocytes. RT-qPCR analysis for the indicated miRNAs is shown as IP/IgG for each independent experiment (IgG arbitrary value 1). Data are the mean± SEM of three independent experiments.

Data are considered statistically significant with p< 0.05 (Student’s T test). *: p<0.05; **: p<0.01; ***: p<0.001.

Furthermore, in line with data obtained for miR-155-3p (see Figure 2D), results shown in Figure 6 demonstrate the SYNCRIP-dependent PCBP2 binding assessed in hepatocytes silenced for SYNCRIP.

PCBP2 binding to miRNAs requires SYNCRIP

CLIP of PCBP2 protein in murine hepatocytes both WT (shCTR) and silenced for SYNCRIP (shSYNCRIP). RT-qPCR analysis for the indicated miRNAs is shown as IP/IgG. Data are the mean± SEM of three independent experiments.

Data are considered statistically significant with p< 0.05 (Student’s T test). *: p<0.05; ***: p<0.001.

Overall, these data indicate that the previously described SYNCRIP capacity to act as EV loader is functionally limited by the presence of a here-identified sequence- and SYNCRIP-dependent retention mechanism mediated by the RNA-binding protein PCBP2.

Discussion

Earlier version: The main finding of this investigation is the identification of PCBP2 as a new regulator of miRNA partition between intracellular and EV compartments. Evidence here gathered indicates that: i) PCBP2 binding requires both the miRNA CELL sequence and the RBP SYNCRIP, which in turn recognizes its specific hEXO consensus; in other words, SYNCRIP binding is a prerequisite for PCBP2 recruitment; ii) PCBP2 cell retention function is dominant over the EV-loading SYNCRIP one; in other word PCBP2 impairs SYNCRIP-mediated miRNA export (Figure 7).

Revised version: The main finding of this investigation is the identification of PCBP2 as a new regulator of miRNA partition between intracellular and EV compartments. Evidence here gathered indicates that: i) PCBP2 binding requires both the miRNA CELL sequence and the RBP SYNCRIP, which in turn recognizes its specific hEXO consensus; in other words, SYNCRIP binding is a prerequisite for PCBP2 recruitment; ii) PCBP2 cell retention function is dominant over the SYNCRIP-dependent EV-loading; in other word PCBP2 impairs SYNCRIP-mediated miRNA export (Figure 7).

Schematic model of PCBP2/SYNCRIP dependent miRNAs compartmentalization.

Left) hEXO-SYNCRIP interaction promotes miRNAs secretion into EVs.

Right) SYNCRIP-dependent PCBP2-CELL motif interaction promotes miRNAs intracellular retention.

While SYNCRIP EV-loading activity has been previously well characterized by means of both functional and structural analysis (7, 9), the role of PCBP2 as mediator of miRNA intracellular retention is here disclosed for the first time.

Indeed, no data were previously reported on the RBPs involved in miRNA cell retention even if the role of specific consensus sequences has been previously defined by means of functional assays involving introduction/removal of the CELL motifs (5).

PCBP2 protein (similarly to SYNCRIP) displays a pleiotropic function; specifically, it is a well-characterized member of the Poly-rC-binding proteins (PCBPs), a group of multifunctional RNA-binding proteins that contain three highly conserved RNA binding KH domains and that may shuttle between the nucleus and the cytoplasm (20). A large body of evidence points to its role in controlling multiple processes including RNA maturation and trafficking, RNA editing, translational activation or repression and mRNA degradation (21–26).

Its potential impact on miRNA biogenesis suggested us to analyze miRNA partition as the EVs/cell ratio thus circumventing variation deriving from the intracellular expression levels.

Here, PCBP2 was found to directly bind to miRNAs, sharing one of the CELL-motifs previously identified by in silico sequence analysis of the miRNAs that were retained in AML12 mouse hepatocytes (5). The use of specific insertion/removal mutants and knock-down cellular systems here highlighted the SYNCRIP recruitment as a prerequisite for PCBP2 binding, this verified on miRs-155-3p, 26b-3p and 31-3p. Furthermore, this conclusion was extended by CLIP analysis to further 9 miRs selected on the basis of an NGS approach and of a bioinformatic analysis.

Of note, among them, miRs-155-3p, 23a-5p, 155-5p, 192-5p and 26b-3p display important EV-mediated functions in relation to pathophysiology (27–30), as also assessed by Gene Ontology analysis.

The here proposed mechanism implies that the export activity of SYNCRIP is specifically impaired by PCBP2 and highlights that the miRNA partition is not only related to the presence of specific RBPs/export sequences interaction. The final functional compartmentalization output appears the result of an integrated system of RNA/proteins interactions, here only partially unveiled, whose dynamics may provide elements for the explanation of the EVs miRNA cargo specificity and for its variation coherently to cellular plasticity.

Earlier version: In addition, it is conceivable to hypothesize that previously identified domains (9, 31–34), shall mediate the protein-protein interaction dynamics between PCBP2 and SYNCRIP, impacting RBP-miRNAs recognition. Moreover, upon the interaction between SYNCRIP and the miRNAs a topological switch may occur, impacting the affinity of PCBP2 for the same miRNAs.

Revised version: In addition, it is conceivable to hypothesize that previously identified domains (9, 31–34), shall mediate the protein-protein interaction dynamics between PCBP2 and SYNCRIP, impacting RBP-miRNAs recognition. Moreover, upon the interaction among SYNCRIP and the miRNAs a topological switch may occur, impacting the affinity of PCBP2 for the same miRNAs.

Of note, recent research highlights a further level of complexity since miRNA epitranscriptomic modifications, while impairing miRNA intracellular function, appear instrumental to miRNA loading in EVs (4).

The described multiple RNA/proteins interactions provide a further step in the process of clarification of the mechanisms that may yield value in the control of cellular communication in pathophysiological processes. Functionally, the effect of PCBP2 silenced cells-derived EVs on the immune response evaluated as IFN-γ production suggests modulation on functional states of recipient target cells. The knowledge of molecular players of miRNAs intracellular/EVs partition could be soon instrumental for the development of RNA-based manipulations holding therapeutic perspectives.

Materials and methods

Cell Culture Conditions

Non tumorigenic murine hepatocyte 3A cells (3, 35) were grown at 37°C, in a humidified atmosphere with 5% CO2, in RPMI 1640 medium supplemented with 10% FBS (Gibco Life Technology), 50 ng/mL epidermal growth factor (EGF), 30 ng/mL insulin growth factor (IGF) II (PeproTech), 10 mg/mL insulin (Roche), and penicillin/streptomycin, on dishes coated with collagen I (Collagen I, Rat Tail; Gibco Life Technology).

Extracellular Vesicle Purification

Extracellular vesicles were prepared according to International Society of Extracellular Vesicles (ISEV) recommendations(36). Conditioned media (CM) from 150 mm plates each containing 250000 hepatocytes were collected after 72-hrs culture in complete medium containing EV-depleted FBS. Cell-conditioned media were centrifuged at 2000 × g for 20 min at 4°C to remove dead cells and then at 20000 × g for 30 min at 4°C. Cleared supernatants were passed through 0.22 mm filter membranes, ultracentrifuged in a SW32 Ti rotor (Beckman Coulter) at 100,000x g for 70 min at 4°C, and finally resuspended in PBS. The EVs resuspension was analyzed by EXOID-V1-SC (IZON) for size and concentration characterization.

Biotin miRNA Pull-Down

Biotin miRNA pull-down experiments were performed on cytoplasmic extracts. Briefly, cells were lysed in hypotonic buffer (10 mM Tris-Cl [pH 7.5], 20 mM KCl, 1.5 mM MgCl2, 5 mM DTT, 0.5 mM EGTA, 5% glycerol, 0.5% NP40, and 40 U/mL RNAsin [Promega]) supplemented with protease inhibitors (Roche Applied Science). Lysates were incubated on a rotating platform for 30 min at 4°C and then centrifuged at 13000 rpm for 30 min at 4°C. Protein concentration was determined with Protein Assay Dye Reagent (Bio-Rad) based on the Bradford assay.

Samples (2 mg of proteins) were incubated for 1 hr at 4°C with 10 nmol synthetic single strand miRNA oligonucleotides containing a biotin modification attached to the 5’ and via a spacer arm (IDT, Intregrated DNA Technology) (Table 1).

Dynabeads™ M-280 Streptavidin (50 μl/sample, Invitrogen™), previously blocked with 1 mg/mL yeast tRNA (Roche Applied Science), were added to reaction mixture for 90 min at 4°C, and then the beads were washed three times with cold lysis buffer and once with PBS. Elution was performed at room temperature for 5 min in Laemli Buffer (containing 2-β mercaptoethanol and SDS).

Detection of miRNA/RBPs interaction was evaluated by WB on 10% of Input sample and 50% of the pulled-down samples.

Pull down assay with PCBP2 (NM_001103165) mouse recombinant protein (TP522190, Origene) was performed with 4 μg of protein.

Protein Digestion, Peptide Purification and nanoLC Analysis

Proteins obtained from the pull-down experiments with miR-155-3p or random scrambled miRNA were separated on 4-12% gradient gels (Invitrogen™) and stained by Simply Blue Safe Stain staining. Fourteen sections of the gel lane were cut. Protein digestion of gel pieces and peptide purification were performed as previously described in (37). Peptides resuspended in a suitable nanoLC injection volume of 2.5% ACN/0.1% TFA and 0.1% formic acid were then analyzed by an UltiMate 3000 RSLCnano-LC system, (Thermo Fisher Scientific) connected on-line via a nano-ESI source to an Q Exactive plus TM Hybrid Quadrupole-OrbitrapTM Mass Spectrometer (Thermo Fisher Scientific) as in (3). Proteins were automatically identified by MaxQuant (v. 1.6.17.0) software. Tandem mass spectra were searched against the Mus Musculus dataset of UniprotKB database.

Quantitative comparison among miR-155-3p WT and miR-155-3p no-CELL was performed using the label-free quantification algorithm calculated by MaxQuant software.

SDS-PAGE and Western Blotting

Cells were lysed in Triton 1X Buffer, subsequently the proteins (30 μg for each sample) were analyzed as in (38). The following primary antibodies were used for immunoblotting: α-PCBP2 (AV40568 – Sigma Aldrich), α-SYNCRIP (MAB11004 – Merck-Millipore), α-HSP90 (sc-13119 - Santa Cruz Biotech.), α-LAMP1 (Ab24170-Abcam), α-CD63 (sc-5275-Santa Cruz Biotechnology), α-SYNTHENIN (Ab133267-Abcam), α-ALIX (2171-Cell Signaling Technology), α-TSG101 (sc-7964, Santa Cruz Biotecnology), α-FLOTILLIN-1 (sc74566-Santa Cruz Biotecnology), α-CALNEXIN (NB100-1965 – Novus Biologicals), α-GAPDH (MAB-374 - Merck-Millipore) used as a loading control. The immune complexes were detected with horseradish peroxidase-conjugated species-specific secondary antiserum: (α-Rabbit 172-1019 and α-Mouse 170-6516 Bio-Rad Laboratories), then by enhanced chemiluminescence reaction (Bio-Rad Laboratories). Densitometric analysis of protein expression was performed by using the Fiji-Image J image processing package.

RNA Extraction, RT-PCR and Real-Time qPCR

miRNAs were extracted by miRNeasy Mini Kit and RNeasy MinElute Cleanup Kit (QIAGEN) and reverse transcribed with MystiCq® microRNA cDNA Synthesis Mix (Sigma-Aldrich). Quantitative polymerase chain reaction (RT-qPCR) analyses were performed according to MIQE guidelines. cDNAs were amplified by qPCR reaction using GoTaq qPCR Master Mix (Promega, Madison, WI, USA). Relative amounts, obtained with 2^(-ΔCt) method, were normalized with respect to the cel-miR-39 Spike-In (59000; NORGEN), previously added into miRNA samples and expressed as ratio of miRNAs expression in EVs with respect to the intracellular compartment (4).

Total RNA was extracted by ReliaPrep™ RNA Tissue Miniprep System (Promega, USA) and reverse transcribed with iScriptTM c-DNA Synthesis Kit (Bio-Rad Laboratories Inc., USA). Quantitative polymerase chain reaction (RT-qPCR) analyses were performed according to MIQE guidelines. cDNAs were amplified by qPCR reaction using GoTaq qPCR Master Mix (Promega, Madison, WI, USA). Relative amounts, obtained with 2^(-ΔCt) method, were normalized with respect to the housekeeping gene 18S. Oligonucleotide sequences are reported in Table 3.

The results were analyzed with Manager Software (Bio-Rad) and calculated by the ΔC(t) method.

Co-Immunoprecipitation

Cells were lysed with IP Lysis Buffer (150 mM NaCl, 50 mM Tris-HCl pH 7.5, 5mM EGTA pH 8, 50 mM NaF pH 8, 1,5 mM MgCl2, 1% TRITON-X100 and 10% glycerol) containing freshly added cocktail protease inhibitors (complete EDTA-free Protease Inhibitor Cocktail; SigmaAldrich) and phosphatase inhibitors (5 mM EGTA pH 8.0; 50 mM sodium fluoride; 5 mM sodium orthovanadate). Lysates were incubated on a rotating platform for 2h at 4°C and then centrifuged at 13000 rpm for 30 min at 4°C. Protein concentration was determined with Protein Assay Dye Reagent (Bio-Rad), based on the Bradford assay.

2 mg of proteins (one for the specific antibody and one for the corresponding aspecific IgG) were precleared adding 40 µL of Protein A Sepharose or Protein G Sepharose (GE HealthCare) for 3 hrs at 4°C in a total volume of 1 ml of IP Lysis Buffer in rotation. Then, Protein A or G Sepharose was removed by centrifugation and the extracts were incubated with 5 µg of specific antibody α-PCBP2 (cod. RN025P - MBL), SYNCRIP (MAB11004, Merck-Millipore), Normal Rabbit IgG (12-370-Millipore) or Normal Mouse IgG (12-371-Millipore), the last two used as negative controls, to proceed with immunoprecipitation at 4°C overnight. Immuno-complexes were collected adding 50 µL of Protein A or G Sepharose for 3 hrs at 4°C in rotation. The immunoprecipitated proteins were washed three times with Net Gel Buffer (150 mM NaCl, 50 mM Tris-HCl pH 7.5, 1mM EDTA, 0.1% NP40 and 0.25% gelatin) and once with RIPA Buffer (150 mM NaCl, 1% NP40, 0,5% Sodium Deoxycholate, 0,1% SDS, 50 mM Tris HCl pH 8). Finally, immunoprecipitated proteins were detached from Sepharose beads by adding 50 µL of Laemli Buffer 2X. Samples were boiled at 95°C for 5 min, beads were eliminated by centrifugation and 10% of input sample and 50% of each immunoprecipitated sample were loaded on polyacrilammide gel and analyzed by Western Blotting.

EMSA

Cells were lysed in Triton Buffer at 4°C, for 30 min and 4 μg of protein extract were incubated with 0.5 pmol of biotinylated RNA oligonucleotides for 30 min at room temperature in REMSA Binding Buffer, according to the manufacturer’s protocol (Light Shift Chemiluminescent RNA EMSA Kit, ThermoScientific 20158). 1 μg of each antibody was incubated with the protein-RNA complex: anti-PCBP2 (RN025P; MBL), anti-SYNCRIP (MAB11004; Merck-Millipore) for supershift and ultrashift analysis. The electrophoresis was performed in native 6% polyacrylamide gel in 0,5X TBE. Transfer step was carried out at 25V, for 15 min in 0,5X TBE and the detection was performed following manufacturer’s instructions.

UV Cross-Linking RIP

CLIP was performed as reported in (38). Immunoprecipitated miRNAs were reverse transcribed and analyzed by RT-qPCR amplifications. List of primers is reported in Table 2. Primary antibodies for IP: anti-PCBP2 (RN025P; MBL), anti-SYNCRIP (MAB11004; Merck-Millipore) and as negative controls Normal Rabbit IgG (12-370; Merck-Millipore) or Normal Mouse IgG (12-371; Merck-Millipore).

shRNA Silencing

Stable PCBP2 knockdown was achieved through infection with shRNAs cloned in pSUPER retro puro retroviral vector (Oligoengine). Viral supernatants were collected 48 hrs after transfection of 293gp packaging cells, filtered (0.45 mm), and added to hepatocytes. At 48 hrs post-infection, selection was performed with 2 µg/mL puromycin for at least 1 week before analysis. The sequence of shRNA scramble used as control was previously described(39). The sequences of shRNA oligos used for cloning are reported in Table 4.

Motif Scanning Analysis

Murine mature miRNA sequences were retrieved from miRBase v22.1 database (40). The FIMO tool (41) was used to scan these sequences for occurrences of hEXO, extended CELL (the bottom motif identified in AML12 cells and reported in figure 2 by (5)) and core AUUA/G CELL motifs, encoded as Position Probability Matrices, with parameters --bfile --motif-- --norc and setting the p-value threshold to 0.1, 0.1 and 0.01, respectively.

Small RNA Sequencing

Earlier version: miRNA samples (two biological replicates per condition), to which the cel-miR-39 Spike-In (59000; NORGEN) was previously added, were sequenced at Procomcure Biotech GmbH. Sequencing libraries were prepared using the NEXTFLEX Small RNA-Seq Kit v4 (PerkinElmer). The sequencing reaction was performed on an llumina NovaSeq 6000 instrument in 2x40bp paired-end configuration, with a throughput of ∼40 million read pairs per sample. FastqToolkit version 2.2.5 (available at https://www.illumina.com/products/by-type/informatics-products/basespace-sequence-hub/apps/fastq-toolkit.html) was used to remove adapter sequences from the 3’ end and to filter out reads whose length and average quality after trimming were < 10 and < 30, respectively. Only forward reads were kept for downstream analyses. The mirPRo software version 1.1.4 (42) which utilizes NovoAlign (43) as its alignment engine, was used to align reads to a reference composed of miRNA hairpin sequences downloaded from miRBase v22.1 database with the addition of the spike-in, and to count reads mapping to mature miRNAs. A count matrix was assembled, including only mature miRNAs with one or more reads in at least two cell and two EV samples. Differential abundance analysis was performed using the DESeq2 R package (44). Size factors were estimated directly from spike-in counts. For each mature miRNA, a likelihood ratio test was conducted to assess differences between the EV/cell abundance ratios measured in the shPCBP2 and shCTR conditions.

Revised version: miRNA samples (two biological replicates per condition), to which the cel-miR-39 Spike-In (59000; NORGEN) was previously added, were sequenced at Procomcure Biotech GmbH. Sequencing libraries were prepared using the NEXTFLEX Small RNA-Seq Kit v4 (PerkinElmer). The sequencing reaction was performed on an llumina NovaSeq 6000 instrument in 2×40bp paired-end configuration, with a throughput of ∼40 million read pairs per sample. FastqToolkit version 2.2.5 (available at https://www.illumina.com/products/by-type/informatics-products/basespace-sequence-hub/apps/fastq-toolkit.html) was used to remove adapter sequences from the 3’ end and to filter out reads whose length and average quality after trimming were < 10 and < 30, respectively. Only forward reads were kept for downstream analyses. The mirPRo software version 1.1.4 (42) which utilizes NovoAlign (43) as its alignment engine, was used to align reads to a reference composed of miRNA hairpin sequences downloaded from miRBase v22.1 database with the addition of the spike-in, and to count reads mapping to mature miRNAs. A count matrix was assembled, including only mature miRNAs with one or more reads in at least two cell and two EV samples. Differential abundance analysis was performed using the DESeq2 R package (44). Size factors were estimated directly from spike-in counts. For each mature miRNA, a likelihood ratio test was conducted to assess differences between the EV/cell abundance ratios measured in the shPCBP2 and shCTR conditions.

miRNA target prediction and enrichment analysis

Enrichment analysis was performed using the R Bioconductor packages ClusterProfiler(v4.14.4), org.Mm.eg.db(v3.20.0) and AnnotationDbi(v1.68.0). The analysis was carried out on miRNA validated targets identified through TarBase v9.0 database (https://dianalab.e-ce.uth.gr/tarbasev9) and miRNA targets predicted with DIANA-MicroT Webserver (https://dianalab.e-ce.uth.gr/microt_webserver/#/), using miRBase v22.1 as miRNA resource. GO terms for Biological Processes were filtered using p-value and q-value cutoffs respectively of 0.05 and 0.2. Multiple testing correction was applied using the Benjamini-Hochberg method. Enriched GO terms were visualized using dotplots.

Ex vivo Analysis

Female C57BL/6J mice were purchased from Charles River and were of 12- to 14-weeks old by the time of use weighing 20–25 g. All animals were housed in ventilated cages (no more than 5 mice per cage) under specific pathogen-free conditions and in a controlled environment (12-h daylight cycle, lights off at 18:00) with free access to food and water. Livers were dissected from mice, smashed using PBS and strained (100-µm filter). Leukocytes, enriched using 40% Percoll, were plated at a concentration of 1×106/ml in complete RPMI, and supplemented with EVs-free 10% Fetal Calf Serum (FCS), 100U/mL penicillin/streptomycin, 1% glutamine and with 20ng/mL recombinant murine IL-2 (Peprotech) and 20ng/mL recombinant murine IL-7 (Peprotech) at 37°C under 5% CO2. Leukocytes were cultured either alone or in the presence of 10 μg/mL control EVs or with the same amount of shPCBP2 EVs for 24 hours at 37°C under 5% CO2. Brefeldin A 5 μg/mL (Merck Life Sciences) was added in the last 3 hrs of treatment.

Cells were stained for cell surface markers using the following antibodies: anti-CD45.2-PercP-Cy5.5 (clone 194, Invitrogen); anti-NK1.1-PE (clone PK136, Invitrogen); anti-CD11b-BV421 (clone M1/70, Biolegend); anti-CD3-BV510 (clone 145-2C11, Biolegend) and APC-H7-conjugated Fixable viability Dye (Thermo Scientific) to identify dead cells. Cells were then fixed and permeabilized using the Cytofix/cytoperm kit (BD Biosciences) before intracellular staining with anti-IFN-γ-APC (clone XMG1.2, Biolegend). Samples were acquired by FacsCantoII and Flow cytometric analysis was performed with FlowJo 10 software (BD Bioscience). Live/dead(L/D)-/CD45.2+ cells were gated and then segregated into NK1.1+ CD11b+, NK1.1- CD11b+, NK1.1- CD11b- CD3+ cells to discriminate NK cells, myeloid cells and T cells, respectively. These gated cells were then analyzed for the expression of IFN-γ as reported in Supplementary Figure 5.

Animal studies were conducted in accordance with all relevant ethical regulations for animal testing and research including the Italian code for the care and use of animals for scientific purposes.

The Italian Ministry of Health approved the use of animals (Authorization n. 698/2021-PR).

Statistical Analyses

For the qRT-PCR analysis, and for ex vivo analysis, statistical differences were assessed with the one-tailed paired Student’s t-test using GraphPad Prism Version 9 (GraphPad Software). Data are presented as mean ± SEM, and p values < 0.05 were considered statistically significant. For the statistical analysis of proteomic studies, Perseus software (version 1.6.7.0) after log2 transformation of the intensity data was used. Results were considered statistically significant at p<0.05.